# ICA7 - Transitioning to Python: Functions

## Write a function to calculate the Pythagorean theorem
$$a^2 + b^2 = c^2$$
1. The function should take 2 parameters (`a` and `b`), compute the value of `c` and return it
2. Call your function with several different values for `a` and `b` and print the results
3. Check your work with a calculator
4. After writing your function in Jupyter Notebook, convert it to a `.py` file and execute it from the command line

Turn in your completed Jupyter Notebook and your `.py` file to GitHub.

<!-- grade_id: cell-6b6160dcf59e988b, points: 0, grade: false, solution: true -->
```python
import math

def pythag(a: int | float, b: int | float) -> float:
    '''Computes the length of the hypotenuse of a right triangle given the lengths
    of the other two sides.'''
    ### BEGIN SOLUTION
    return math.sqrt(a**2 + b**2)
    ### END SOLUTION
```

```python
# Use this cell to check several outputs of your function
pythag(3,4)
```

Output:
```
5.0
```

<!-- grade_id: cell-306aa06c9d6b0fd9, points: 5, grade: true, solution: false -->
```python
'''Check function results, do not change this cell'''
assert pythag(3,4) == 5, "fails 3,5 test"
assert pythag(5,12) == 13, "fails 5,12 test"
assert pythag(9,40) == 41, "fails 9,40 test"
### BEGIN HIDDEN TESTS
assert pythag(13,84) == 85, "fails hidden test1"
assert pythag(65,72) == 97, "fails hidden test2"
### END HIDDEN TESTS
print("You pass the tests! \U0001f600")
```

Output:
```
You pass the tests! 😀
```

# Convert your Jupyter notebook to a script

1. Depending on which version of Jupyter Notebook you're running either
    - Select File $\rightarrow$ Download As $\rightarrow$ Python (.py) or
    - File $\rightarrow$ Save and Export Notebook As... $\rightarrow$ Executable Script
3. Open your ```.py``` file in VS Code and use the pylance to check the type annotations. (Hopefully they're all correct, since we wrote them 😝 )
4. Add and upload `ICA7.py` to your repo. 
5. Include the name of your script file in the text box below.
6. REMEMBER TO PUSH ALL CHANGES TO YOUR ONLINE REPO.

<!-- grade_id: cell-3d3180aeed96e443, points: 5, grade: true, solution: true -->
