# Bi 621 – Problem Set 3

The goal of this problem set is to convert phred quality scores from their single letter encoding into integers and to calculate the average phred quality score for a single sequence. **Be sure to remove the ```raise NotImplementedError()``` from the code.**

```python
#initializing variables, be sure to run this cell and do not change
seq_name: str = "@GAAATG_1_1101_4446_2137_1"
seq: str = "TGCAGGTTGAGTTCTCGCTGTCCCGCCCCCATCTCTTCTCTTCCAGTCTGGCTCTGGAGCAGTTGAGCCCAGCTCAGGTCCGCCGGAGGAGACCG"
phred_score: str = "FFHHHHHJJJJIJIJJJIJJJJJJIIIJJJEHJJJJJJJIJIDGEHIJJFIGGGHFGHGFFF@EEDE@C??DDDDDDD@CDDDDBBDDDBDBDD@"
```

# Part 1 – Convert_phred
Write a function that takes a single character as it's parameter and converts that parameter into a phred score, and returns the phred score.

>*Remember: Illumina 1.8+ phred scores are offset by a factor of 33 from 0. Review the lecture notes as well as the [Wikipedia FASTQ page](https://en.wikipedia.org/wiki/FASTQ_format#Encoding) to see how to do the conversion.*

*HINT* – what does the `ord()` function do? Check the documentation

### REMEBER to add `qual_score()` to the `Bioinfo.py` Python module you are creating. 

<!-- grade_id: cell-be6acddb59ad29cd, points: 0, grade: false, solution: true -->
```python
def convert_phred(letter: str) -> int:
    """Converts a single character into a phred score"""
    ### BEGIN SOLUTION
    return ord(letter)-33
    ### END SOLUTION
```

```python
# Check to see that your code works
# Recommend swapping out "A" for several different ASCII values to check a variety of test cases.
convert_phred('A')
```

Output:
```
32
```

<!-- grade_id: cell-3df715c9006e3932, points: 3, grade: true, solution: false -->
```python
#autograder tests, do not change
"""Check that convert_phred returns the correct value for several different inputs"""
assert convert_phred("I") == 40, "wrong phred score for 'I'"
assert convert_phred("C") == 34, "wrong phred score for 'C'"
assert convert_phred("2") == 17, "wrong phred score for '2'"
assert convert_phred("@") == 31, "wrong phred score for '@'"
assert convert_phred("$") == 3, "wrong phred score for '$'"
print("Your convert_phred function is working! Nice job")
```

Output:
```
Your convert_phred function is working! Nice job
```

# Part 2 – Convert the entire phred string
Write a loop print each letter in the `phred_score` string and its corresponding score in the **following format**:
```
0: F - 37
1: F - 37
2: H - 39
3: H - 39
4: H - 39
5: H - 39
6: H - 39
7: J - 41
...
```

Reminder, start simply.
* Write a loop to iterate over each letter in the `phred_score` and print them out.
* Modify your loop to print the letter and the corresponding score. *Be sure to use your `convert_phred` function*.
* Alter the loop to print out an index: 0 for the first letter, 1 for the 2nd letter, etc.
* **Be sure your formatting matches the requested output**.

<!-- grade_id: cell-07c9f3140778719b, points: 3, grade: true, solution: true -->
```python
### BEGIN SOLUTION
for i in range(len(phred_score)):
    #print(str(i) + ': ' + str(phred_score[i]) + ' - ' + str(convert_phred(phred_score[i])))
    print(f"{i}: {phred_score[i]} - {convert_phred(phred_score[i])}")
### END SOLUTION
```

Output:
```
0: F - 37
1: F - 37
2: H - 39
3: H - 39
4: H - 39
5: H - 39
6: H - 39
7: J - 41
8: J - 41
9: J - 41
10: J - 41
11: I - 40
12: J - 41
13: I - 40
14: J - 41
15: J - 41
16: J - 41
17: I - 40
18: J - 41
19: J - 41
20: J - 41
21: J - 41
22: J - 41
23: J - 41
24: I - 40
25: I - 40
26: I - 40
27: J - 41
28: J - 41
29: J - 41
30: E - 36
31: H - 39
32: J - 41
33: J - 41
34: J - 41
35: J - 41
36: J - 41
37: J - 41
38: J - 41
39: I - 40
40: J - 41
41: I - 40
42: D - 35
43: G - 38
44: E - 36
45: H - 39
46: I - 40
47: J - 41
48: J - 41
49: F - 37
50: I - 40
51: G - 38
52: G - 38
53: G - 38
54: H - 39
55: F - 37
56: G - 38
57: H - 39
58: G - 38
59: F - 37
60: F - 37
61: F - 37
62: @ - 31
63: E - 36
64: E - 36
65: D - 35
66: E - 36
67: @ - 31
68: C - 34
69: ? - 30
70: ? - 30
71: D - 35
72: D - 35
73: D - 35
74: D - 35
75: D - 35
76: D - 35
77: D - 35
78: @ - 31
79: C - 34
80: D - 35
81: D - 35
82: D - 35
83: D - 35
84: B - 33
85: B - 33
86: D - 35
87: D - 35
88: D - 35
89: B - 33
90: D - 35
91: B - 3
```

Write a second function that takes the original, unmodified `phred_score` string as a parameter. This function should calculate the average quality score of the whole phred string. Be sure to write a doc string for this funcion.
> Struggling? See some [hints](PS3_hints.ipynb)

### REMEBER to add `qual_score()` to the `Bioinfo.py` Python module you are creating. 

<!-- grade_id: cell-3d1f2021427e7b65, points: 0, grade: false, solution: true -->
```python
def qual_score(phred_score: str) -> float:
    """Write your own doc string"""
    ### BEGIN SOLUTION
    sum: int = 0
    for item in phred_score:
        sum += convert_phred(item)
    return sum/len(phred_score)
    ### END SOLUTION
```

```python
#Use this cell to call and test your qual_score() function
qual_score(phred_score)
```

Output:
```
37.62105263157895
```

<!-- grade_id: cell-a287d039bcc62114, points: 1, grade: true, solution: false -->
```python
#autograder tests, do not change
old_convert_phred = convert_phred
del convert_phred

try:
    qual_score(phred_score)
except NameError:
    pass
else:
    raise AssertionError("qual_score does not call convert_phred - make sure you're not duplicating previous effort!")
finally:
    convert_phred = old_convert_phred
    del old_convert_phred
    
print("you used the function you wrote, excellent!")
```

Output:
```
you used the function you wrote, excellent!
```

<!-- grade_id: cell-9e105f0f82dea141, points: 3, grade: true, solution: false -->
```python
#autograder tests, do not change
assert qual_score("EEE") == 36
assert qual_score("#I") == 21
assert qual_score("EJ") == 38.5
assert qual_score(phred_score) == 37.62105263157895, "wrong average phred score"
print("You calcluated the correct average phred score")
```

Output:
```
You calcluated the correct average phred score
```

```python

```